```python
# Initialize Otter
import otter
grader = otter.Notebook("Assignment0.ipynb")
```

# CMPUT 200 Winter 2022
# Assignment 0

Each assignment will be distributed as a notebook such as this one.  You will execute the questions in the notebook.  The questions might ask for a short answer in text form or for you to write and execute a piece of code. 
For short answer questions you must enter your answer in the provided space.  For coding questions you must use the provided space.  When you are done, you will submit your work from the notebook. Follow directions at the bottom of this notebook for submission.

```python
# Don't change this cell; just run it. 

%pip install -r requirements.txt
import numpy as np
import pandas as pd
from scipy.optimize import minimize

# These lines do some fancy plotting magic.
import matplotlib
# This is a magic function that renders the figure in the notebook, instead of displaying a dump of the figure object.
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)

import otter
grader = otter.Notebook()
```

## Importing Data

## Linear Regression

```python
# Just run this cell
pten = pd.read_csv("pten.csv")
pten.head(5)
```

```python
# Just run this cell
pten.hist("Copy Number", bins = np.arange(-1, 1.5, 0.5))
```

**Question 1**

Looking at the histogram above, we want to check whether or not `Copy Number` is in standard units. For this question, compute the mean and the standard deviation of the values in `Copy Number` and assign these values to `copy_number_mean` and `copy_number_sd` respectively. After you calculate these values, assign `is_su` to either `True` if you think that `Copy Numbers` is in standard units or `False` if you think otherwise.

```python
copy_number = pten['Copy Number']
copy_number_mean = ...
copy_number_sd = ...
is_su = ...
print(f"Mean: {copy_number_mean}, SD: {copy_number_sd}, Is in standard units?: {is_su}")
```

```python
grader.check("q1")
```

**Question 2**

Create the function `standard_units` so that it converts the values in the array `arr` to standard units. We'll then use `standard_units` to create a new table, `pten_su`, that converts all the values in the table `pten` to standard units.

```python
def standard_units(arr):
    ...

# DON'T DELETE OR MODIFY ANY OF THE LINES OF CODE BELOW IN THIS CELL
pten_su = pd.DataFrame({"Cell Line": pten["Cell Line"],
            "Copy Number SU": standard_units(pten["Copy Number"]),
            "mRNA Expression (Affy) SU": standard_units(pten["mRNA Expression (Affy)"]),
            "mRNA Expression (RNAseq) SU": standard_units(pten["mRNA Expression (RNAseq)"])}) 
print(pten_su.head(5))
```

```python
grader.check("q2")
```

You should always visually inspect your data, before numerically analyzing any relationships in your dataset. Run the following cell in order to look at the relationship between the variables in our dataset.

```python
# Just run this cell
pten_su.plot.scatter("Copy Number SU", "mRNA Expression (Affy) SU")
pten_su.plot.scatter("Copy Number SU", "mRNA Expression (RNAseq) SU")
pten_su.plot.scatter("mRNA Expression (Affy) SU", "mRNA Expression (RNAseq) SU")
```

**Question 3**

Which of the following relationships do you think has the highest correlation (i.e. highest absolute value of `r`)? Assign `highest_correlation` to the number corresponding to the relationship you think has the highest correlation.

1. Copy Number vs. mRNA Expression (Affy)
2. Copy Number vs. mRNA Expression (RNAseq)
3. mRNA Expression (Affy) vs. mRNA Expression (RNAseq)

```python
highest_correlation = ...
```

**Question 4**

Now, using the `standard units` function, define the function `correlation` which computes the correlation between `arr1` and `arr2`.

```python
def correlation(arr1, arr2):
    ...
# This computes the correlation between the different variables in pten
copy_affy = correlation(pten["Copy Number"], pten["mRNA Expression (Affy)"])
copy_rnaseq = correlation(pten["Copy Number"], pten["mRNA Expression (RNAseq)"])
affy_rnaseq = correlation(pten["mRNA Expression (Affy)"], pten["mRNA Expression (RNAseq)"])

print(f" \
      Copy Number vs. mRNA Expression (Affy) Correlation: {copy_affy}, \n \
      Copy Number vs. mRNA Expression (RNAseq) Correlation: {copy_rnaseq}, \n \
      mRNA Expression (Affy) vs. mRNA Expression (RNAseq) Correlation: {affy_rnaseq}")
```

```python
grader.check("q4")
```

**Question 5**

If we switch what we input as arguments to `correlation`, i.e. found the correlation between `mRNA Expression (Affy)` vs. `Copy Number` instead of the other way around, would the correlation change? Assign `correlation_change` to either `True` if you think yes, or `False` if you think no.

```python
correlation_change = ...
```

<!-- BEGIN QUESTION -->

**Question 6**

Looking at both the scatter plots after Question 2 and the correlations computed in Question 4, describe a pattern you see in the relationships between the variables.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 7**

Let's look at the relationship between mRNA Expression (Affy) vs. mRNA Expression (RNAseq) only. Define a function called `regression_parameters` that returns the parameters of the regression line as a two-item array containing the slope and intercept of the regression line as the first and second elements respectively. The function `regression_line` takes in two arguments, an array of `x` values, and an array of `y` values.

```python
def regression_parameters(x, y):
    ...
    slope = ...
    intercept = ...
    return [slope, intercept]
    
parameters = regression_parameters(pten["mRNA Expression (Affy)"], pten["mRNA Expression (RNAseq)"])
parameters
```

```python
grader.check("q7")
```

**Question 8**

If we switch what we input as arguments to `regression_parameters`, i.e. found the parameters for the regression line for `mRNA Expression (RNAseq)` vs. `mRNA Expression (Affy)` instead of the other way around, would the regression parameters change (would the slope and/or intercept change)? Assign `parameters_change` to either `True` if you think yes, or `False` if you think no.

```python
parameters_change = ...
```

**Question 9**

Now, let's look at how the regression parameters look like in standard units. Use the table `pten_su` and the function `regression_parameters`, and assign `parameters_su` to a two-item array containing the slope and the intercept of the regression line for mRNA Expression (Affy) in standard units vs. mRNA Expression (RNAseq) in standard units.

```python
parameters_su = ...
parameters_su
```

```python
grader.check("q9")
```

<!-- BEGIN QUESTION -->

**Question 10**

Looking at the array `parameters_su`, what do you notice about the slope and intercept values specifically? Relate them to another value we already calculated in a previous question, as well as relate them to an equation.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 11**

The oldest and most commonly used cell line in Biology is the HeLa cell line, named after Henrietta Lacks, whose cervical cancer cells were taken without her consent in 1951 to create this cell line. The issue of data privacy and consent is very important to data science! You can read more about this topic [here](https://www.hopkinsmedicine.org/henriettalacks/).

The HeLa cell line is missing from our dataset. If we know that the HeLa mRNA Expression (Affy) value is 8.2, what is the predicted mRNA Expression (RNAseq) value? Use the values in `parameters` that we derived in Question 7, and assign the result to `hela_rnaseq`.

```python
hela_rnaseq = ...
```

```python
grader.check("q11")
```

**Question 12**

Compute the predicted mRNA Expression (RNAseq) values from the mRNA Expression (Affy) values in the `pten` table. Use the values in the `parameters` array from Question 7, and assign the result to `predicted_rnaseq`. We'll plot your computed regression line with the scatter plot from after question 2 of mRNA Expression (Affy) vs. mRNA Expression (RNAseq).

```python
predicted_rnaseq = ...
# DON'T CHANGE/DELETE ANY OF THE BELOW CODE IN THIS CELL
pten["Predicted mRNA Expression (RNAseq)"] = predicted_rnaseq
pten[["mRNA Expression (Affy)", "mRNA Expression (RNAseq)", "Predicted mRNA Expression (RNAseq)"]].plot.scatter("mRNA Expression (Affy)", "mRNA Expression (RNAseq)")
plt.plot(pten["mRNA Expression (Affy)"], predicted_rnaseq)
```

## Fitting a least-squares regression line

Recall that the least-square regression line is the unique straight line that minimizes root mean squared error (RMSE) among all possible fit lines. Using this property, we can find the equation of the regression line by finding the pair of slope and intercept values that minimize root mean squared error.

**Question 13**

Define a function called `RMSE`. It takes in one argument 'params' which is a two-item array. The items are:

1. the slope of a line (a number)
2. the intercept of a line (a number).

It should return a number that is the root mean squared error (RMSE) for a line defined with the arguments slope and intercept used to predict mRNA Expression (RNAseq) values from mRNA Expression (Affy) values for each row in the `pten` table.

*Hint: Errors are defined as the difference between the actual `y` values and the predicted `y` values.*

*Note: if you need a refresher on RMSE, here's the [link](https://www.inferentialthinking.com/chapters/15/3/Method_of_Least_Squares.html#Root-Mean-Squared-Error) from the textbook*

```python
def RMSE(params):
    slope, intercept = params[0], params[1]
    affy = pten["mRNA Expression (Affy)"]
    rnaseq = pten["mRNA Expression (RNAseq)"]
    predicted_rnaseq = ...
    ...

# DON'T CHANGE THE FOLLOWING LINES BELOW IN THIS CELL
rmse_example = RMSE([0.5, 6])
rmse_example
```

```python
grader.check("q13")
```

<!-- BEGIN QUESTION -->

**Question 14**

What is the RMSE of a line with slope 0 and intercept of the mean of `y` equal to?

*Hint 1: The line with slope 0 and intercept of mean of `y` is just a straight horizontal line at the mean of `y`*

*Hint 2: What does the formula for RMSE become if we input our predicted `y` values in the formula. Try writing it out on paper! It should be a familiar formula.*

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 15**

Find the parameters that minimizes RMSE of the regression line for mRNA Expression (Affy) vs. mRNA Expression (RNAseq). Assign the result to `minimized_parameters`.

You will have to use the `minimize` [function](https://docs.scipy.org/doc/scipy/reference/generated/scipy.optimize.minimize.html) from the scipy library.

*Hint: Use the `RMSE` function in Question 13*

```python
minimized_parameters = ...
minimized_parameters.x
```

```python
grader.check("q15")
```

<!-- BEGIN QUESTION -->

**Question 16**

The slope and intercept pair you found in Question 15 should be very similar to the values that you found in Question 7. Why were we able to minimize RMSE to find the same slope and intercept from the previous formulas?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 17**

If we had instead minimized mean squared error (MSE), would we have gotten the same slope and intercept of the minimized root mean squared error (RMSE) results? Assign `same_parameters` to either `True` if you think yes, or `False` if you think no.

_Type your answer here, replacing this text._

```python
same_parameters = ...
same_parameters
```

```python
grader.check("q17")
```

<!-- BEGIN QUESTION -->

**Question 18**

Using a linear regression model, would we be able to obtain accurate predictions for most of the points? Explain why or why not.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Convert manually graded questions to pdf**

Running the following cell will convert the manually graded questions to pdf. 

```python
! otter export -e html "Assignment0.ipynb" --filtering
```

## SUBMISSION INSTRUCTIONS
This is the end of Assignment 0.  Be sure to run the tests and verify that they all pass (just because the tests pass does not mean it's the right answer). For submission you need to submit *2 files* on eclass: 

1- a zip file "CCID.zip", this zip file will only include this notebook  
2- pdf of the manually graded questions

This assignment contains altered snippets from the original [Berkeley data-8 course](http://data8.org/), which is licensed under the [Creative Commons license](https://creativecommons.org/licenses/by-nc/4.0/).

Tests for q1 (points: None)
```
>>> np.isclose(copy_number_mean, -0.19447913669064748)
True
```
```
>>> np.isclose(copy_number_sd, 0.7466239639509008)
True
```


Tests for q11 (points: None)
```
>>> np.isclose(hela_rnaseq, 2.329061)
True
```


Tests for q13 (points: None)
```
>>> assert np.isclose(rmse_example, 7.612008179226994)
```


Tests for q15 (points: None)
```
>>> assert np.allclose(minimized_parameters.x, [1.0711393, -6.454280])
```


Tests for q17 (points: None)


Tests for q2 (points: None)
```
>>> assert np.allclose(standard_units(pten["Copy Number"][0:10]), [0.828166, 0.849192, 0.354493, 0.613627, -0.643803, -1.830297, 0.531138, 0.641585, 0.496132, -1.840233])
```


Tests for q3 (points: None)


Tests for q4 (points: None)
```
>>> assert np.isclose(copy_affy,0.581951)
```


Tests for q5 (points: None)


Tests for q7 (points: None)
```
>>> assert np.allclose(regression_parameters(pten["mRNA Expression (Affy)"], pten["mRNA Expression (RNAseq)"]), [1.0711396, -6.454283])
```


Tests for q8 (points: None)


Tests for q9 (points: None)
```
>>> assert np.allclose(parameters_su, [0.900076, 1.648472e-16])
```
